package main

import "fmt"

// Make it compile!
// In Go there are multiple ways to declare a variables.
// I AM NOT DONE

var myFloat

func main() {
	// Declare myInt with type int
    myInt:=
	fmt.Printf("Value of variable myInt of type int is %d", myInt)

	// Declare myString with type string
    myString
	fmt.Printf("Value y of type string is %s", myString)

	// Declare myFloat with type int
	fmt.Printf("Value myFloat of type float is %.2f", myFloat)
}
