package main

import "fmt"

// Make it compile!
// Multiple variables can also be declared together in one variable declaration.

// I AM NOT DONE
func main() {
	// Complete the line without changing anything.
	var hello, world string =

	//print out hello,world
	fmt.Printf("%s,%s", hello, world)

	var (
		one, two, three = "one"
		four            = 4
	)
	fmt.Printf("%s,%s,%s,%d", one, two, three, four)
}
