package packages5

// Make it compile!
// The default name of the package declaration is the name of the package. However it can be override with an identifier.

// I AM NOT DONE

import (
	"fmt"

	"github.com/jeffreylean/gopherize/exercise/packages/lib"
)

// Don't modify the main function.
func main() {
	total := library.Sum(1, 1)
	fmt.Println("Sum of 1 + 1 =", total)
}
