package main

// Make it compile!

// I AM NOT DONE
func main() {
	nums := []int{2, 3, 4}
	sum := 0

    // add the range keyword here to iterate through the nums array, and perform addition for each value in the array and assign it to sum variable.
	for _, num:=  nums {

	}

    fmt.Printf("The sum of an array is %v",sum)
}
