package for2

// Make the test pass!
// The elegance of Go lies in its simplicity. Go provides a single way to create loops, which is the for syntax.
// There is no separate while syntax for creating while loops. Instead, Go uses the versatile for syntax to cover all types of loops.
// In this exercise, your task is to implement while loops using the for syntax in Go.

// I AM NOT DONE
func loop() int {
	i := 0

	// Modify the function so that the value of variable i is 10 by applying while loop using `for` syntax in this block.
	{
		i++
	}

	return i
}
